# <div style="margin: 1px; background-color: #555; color: white; padding:8px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">CNG 465 • Spring 2024-2025 • <span style="margin-left: 1px; background-color: #cf441e; border-radius: 3px; padding: 2px; padding-left: 5px; padding-right: 5px; font-size: 80%;">HOMEWORK 01</span></div>

### <span style="display: block; margin-left: 0px; background-color: #388; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 0px; text-align: center; border-radius: 3px;">3 questions / 100 points total</span><span id="QTOP"></span> 

Jump to 
<a href="#Q1" style="background-color: #555; color: white; padding: 10px; border-radius: 2px; text-align: center; border-radius: 3px; margin-right: 5px;">Q1</a>
<a href="#Q2" style="background-color: #555; color: white; padding: 10px; border-radius: 2px; text-align: center; border-radius: 3px; margin-right: 5px;">Q2</a>
<a href="#Q3" style="background-color: #555; color: white; padding: 10px; border-radius: 2px; text-align: center; border-radius: 3px; margin-right: 5px;">Q3</a>

<h3><span style="background-color: #444; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; margin-top: 10px; margin-bottom: 10px; border-radius: 7px;">INSTRUCTIONS</span></h3>

<div style="color: #515; font-size: 12pt; line-height: 1.5; margin-top: 20px;">

- <b>Note that downloading this notebook, reading and comprehending all instructions contained therein, reading and comprehending the questions on this assignment, filling in/saving your answers in the appropriate cells in this file, and finally uploading this file with your answers onto the appropriate website before the deadline are all parts of this assignment!</b>
    
    
- No time extension will be provided unless a system-related problem occurs!    
    
    
- All questions in this assignment are subject to being automatically graded using [NBGRADER](https://nbgrader.readthedocs.io/).
    
    
- <div style="background-color: #800; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; margin-top: 10px; margin-bottom: 10px; border-radius: 7px;">Yet your answers may selectively be checked for <b>plagiarism</b> using specialized software. If there is clear evidence of plagiarism anywhere in this assignment, you will automatically receive no grade for this assignment! Even if you plagiarize on a single question or one part of a question, you will still receive a <b>zero<b> for the entire assignment, and further disciplinary action may also be taken. So, <b>please do not plagiarize!<b></div>
    
- __<span style="color: #b00; font-size: 12pt; line-height: 1.5;">DO NOT COPY AND PASTE CELLS</span>__
    
- __<span style="color: #b00; font-size: 12pt; line-height: 1.5;">DO NOT REMOVE CELLS</span>__
    
- __<span style="color: #b00; font-size: 12pt; line-height: 1.5;">DO NOT MOVE THE LOCATION OF ANY CELLS</span>__    

- __<span style="color: #b00; font-size: 12pt; line-height: 1.5;">DO NOT ADD NEW CELLS</span>__    

- __<span style="color: #b00; font-size: 12pt; line-height: 1.5;">DO NOT CHANGE THE TYPE (Markdown/Code/Raw) OF CELLS</span>__    
    
- __<span style="color: #080; font-size: 12pt; line-height: 1.5;">ONLY WRITE CODE IN DESIGNATED CELLS AND RUN TESTS</span>__    

- Fill in the code cells indicated for you clearly. That is, do __not__ write your code anywhere else. If you do so, you may lose the points for that question.
    
- Do not wait until the very last moment to upload your work!
    
- Do not send a file with a different file type&mdash;Only Jupyter notebooks will be accepted
    
- Do not send your answer notebook file by email or via any other messaging media&mdash;such requests will be ignored without notice
    
- __All _Rules of Academic Conduct_ apply__. Any sort of cheating will result in an __FF__ grade for the course. Your case will also be submitted to the Faculty Disciplinary Committee for further evaluation and disciplinary action. You can refer to books and web pages for help, but you are __not__ allowed to consult anyone online. Any sort of collaboraton with anyone anywhere is considered cheating.
    
- Enter your official name and your METU NCC student ID in the code cell further below, *unless your name and student ID have already been entered for you*. For example, the format of your information should be as follows:

        student_name = "John Smith"
        student_id   = "2099234"
        
    - Make sure you enter your official name as a `str` as shown above
    - Make sure that your student ID has been entered correctly as `str`.

<h3><span style="background-color: #444; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; margin-top: 10px; margin-bottom: 10px; border-radius: 7px;">SUBMISSION</span></h3>

- Ideally, do **not** modify the name of your notebook file

- Before uploading your work, you must **_validate_** your notebook file using [NBGRADER](https://nbgrader.readthedocs.io/en/stable/)

- Follow the instructions [here](https://nbgrader.readthedocs.io/en/stable/) to install and use [NBGRADER](https://nbgrader.readthedocs.io/en/stable/) for validation

- If your notebook fails to *validate*, you must fix your file before submitting your file!

- Notebooks that do not validate cannot be graded! Therefore, if you upload your file without *validating* it, you run the risk of getting **zero** for your submission! We will **not** take the time to fix your file for you!
    
- Upload your notebook file with your answers at the appropriate ODTUClass link **after validating it using NBGRADER** (*See above*)

- You may upload multiple times, but always ensure that you have removed any previous versions!

- No late submissions will be accepted!

### Best of luck!

</div>

<div style="border-bottom: 2px solid;"></div>

### <span style="margin-left: 0px; background-color: #116; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px; font-size: 12pt;">Check to make sure that your name and student ID are correct in the code cell below</span>

```python
student_name = "Iclal Akman"
student_id   = "2584605"
```

<div style="border-bottom: 2px solid;"></div>

### Computation of *complementary*/*reverse complementary*/*reverse* DNA sequences

From https://study.com/skill/learn/writing-complementary-dna-sequences-explanation.html

**Example Problem 1 - Writing Complementary DNA Sequences:**

What is the complementary sequence for the DNA strand `ATAACGTGGATGCTA`?

**Step 1:** Starting on the left, read the first letter in the DNA sequence and write down the complementary base.

The first base in this sequence is "A", which pairs with "T", so the first letter in our complementary sequence will be "T".

**Step 2:** Continue from left to right until the complementary sequence is complete.

Working through the sequence, the complementary strand will be: `TATTGCACCTACGAT`

**Example Problem 2 - Writing Complementary DNA Sequences:**

What is the complementary sequence for the DNA strand `GTCATGCTAAGTCTGA`?

**Step 1:** Starting on the left, read the first letter in the DNA sequence and write down the complementary base.

The first base in this sequence is "G", which pairs with "C", so the first letter in our complementary sequence will be "C".

**Step 2:** Continue from left to right until the complementary sequence is complete.

Working through the sequence, the complementary strand will be: `CAGTACGATTCAGACT`

---

To test your work using the following online application:

> https://www.bioinformatics.org/sms/rev_comp.html

For *complements*, make sure that the bottom right drop-down menu shows "complement." If not, you must select it. For *reverse complement*, the menu option must be "reverse complement." Finally, for *reverse*, the menu option must be "reverse."

### <span id="Q1" style="background-color: #555; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">Question 1</span><span style="margin-left: 10px; background-color: #383; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">30 points</span><a href="#QTOP" style="background-color: #aaa; color: white; padding: 5px; padding-bottom: 5px; border-radius: 2px; text-align: center; border-radius: 3px; margin-left: 10px; text-decoration: none;">⇡</a>

Write a Python function called `dna_complement( sequence )` that takes a DNA sequence as input and returns the left-to-right complementary sequence for that sequence as a string. 

- The input sequence may be in uppercase, lowercase, or a mixture, but you may assume the input string will always contain bases and nothing else, except for spaces around the the string.

- An empty function stub is provided for you in the answer code cell below.

<span style="margin-left: 0px; background-color: #808; color: white; padding: 5px; margin: 10px; border-radius: 5px; font-size: 12pt;"><b>IMPLEMENT YOUR ANSWER IN THE FOLLOWING CODE CELL</b></span>

<!-- grade_id: cell-r8217t8n04842t7i, points: 0, grade: false, solution: true -->
```python
def dna_complement( sequence ):
    result = ""
    
    for char in sequence:
        if char == "A" or char == "a":
            result += "T"
        elif char == "T" or char == "t":
            result += "A"
        elif char == "G" or char == "g":
            result += "C"
        elif char == "C" or char == "c":
            result += "G"

    return result    
    
        
```

### <span id="Q1p3" style="background-color: #515; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">TEST 1/2 (15 points)</span><a href="#QTOP" style="background-color: #aaa; color: white; padding: 5px; padding-bottom: 5px; border-radius: 2px; text-align: center; border-radius: 3px; margin-left: 10px; text-decoration: none;">⇡</a>

<!-- grade_id: cell-otzs399l1022q3bs, points: 15, grade: true, solution: false -->
```python
sequence = 'ATAACGTGGATGCTA'

student_answer = dna_complement( sequence )
correct_answer = 'TATTGCACCTACGAT'

assert student_answer == correct_answer
print( "+++ CORRECT!" )

```

Output:
```
+++ CORRECT!
```

### <span id="Q1p3" style="background-color: #515; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">TEST 2/2 (15 points)</span><a href="#QTOP" style="background-color: #aaa; color: white; padding: 5px; padding-bottom: 5px; border-radius: 2px; text-align: center; border-radius: 3px; margin-left: 10px; text-decoration: none;">⇡</a>

<!-- grade_id: cell-tcatd82pn0las9o0, points: 15, grade: true, solution: false -->
```python
sequence = 'GTCATGCTAAGTCTGA'

student_answer = dna_complement( sequence )
correct_answer = 'CAGTACGATTCAGACT'

assert student_answer == correct_answer
print( "+++ CORRECT!" )


```

Output:
```
+++ CORRECT!
```

### <span id="Q2" style="background-color: #555; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">Question 2</span><span style="margin-left: 10px; background-color: #383; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">30 points</span><a href="#QTOP" style="background-color: #aaa; color: white; padding: 5px; padding-bottom: 5px; border-radius: 2px; text-align: center; border-radius: 3px; margin-left: 10px; text-decoration: none;">⇡</a>

Write a Python function called `dna_reverse_complement( sequence )` that takes a DNA sequence as input and returns the reverse complementary sequence for that sequence as a string read from right-to-left. This is the method described in the *UBI* textbook where the "complementary" sequence of `AAG` is given as `CTT`. 

- The input sequence may be in uppercase, lowercase, or a mixture, but you may assume the input string will always contain bases and nothing else, except for spaces around the the string.

- An empty function stub is provided for you in the answer code cell below.

<span style="margin-left: 0px; background-color: #808; color: white; padding: 5px; margin: 10px; border-radius: 5px; font-size: 12pt;"><b>IMPLEMENT YOUR ANSWER IN THE FOLLOWING CODE CELL</b></span>

<!-- grade_id: cell-2m48r00p0j21af07, points: 0, grade: false, solution: true -->
```python
def dna_reverse_complement( sequence ):
    result = dna_complement(sequence)
    return result[::-1]
        
    
```

### <span id="Q1p3" style="background-color: #515; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">TEST 1/2 (15 points)</span><a href="#QTOP" style="background-color: #aaa; color: white; padding: 5px; padding-bottom: 5px; border-radius: 2px; text-align: center; border-radius: 3px; margin-left: 10px; text-decoration: none;">⇡</a>

<!-- grade_id: cell-8jh1o47uoq1wyp5m, points: 15, grade: true, solution: false -->
```python
sequence = 'ATAACGTGGATGCTA'

student_answer = dna_reverse_complement( sequence )
correct_answer = 'TAGCATCCACGTTAT'

assert student_answer == correct_answer
print( "+++ CORRECT!" )

```

Output:
```
+++ CORRECT!
```

### <span id="Q1p3" style="background-color: #515; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">TEST 2/2 (15 points)</span><a href="#QTOP" style="background-color: #aaa; color: white; padding: 5px; padding-bottom: 5px; border-radius: 2px; text-align: center; border-radius: 3px; margin-left: 10px; text-decoration: none;">⇡</a>

<!-- grade_id: cell-970336e2d9f18f99, points: 15, grade: true, solution: false -->
```python
sequence = 'GTCATGCTAAGTCTGA'

student_answer = dna_reverse_complement( sequence )
correct_answer = 'TCAGACTTAGCATGAC'

assert student_answer == correct_answer
print( "+++ CORRECT!" )


```

Output:
```
+++ CORRECT!
```

### <span id="Q3" style="background-color: #555; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">Question 3</span><span style="margin-left: 10px; background-color: #383; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">40 points</span><a href="#QTOP" style="background-color: #aaa; color: white; padding: 5px; padding-bottom: 5px; border-radius: 2px; text-align: center; border-radius: 3px; margin-left: 10px; text-decoration: none;">⇡</a>

Write a Python function called `dna_reverse( sequence )` that takes a DNA sequence as input and returns the reverse of that sequence as a string. No base name translation is done in this case!

- The input sequence may be in uppercase, lowercase, or a mixture, but you may assume the input string will always contain bases and nothing else, except for spaces around the the string.

- An empty function stub is provided for you in the answer code cell below.

<span style="margin-left: 0px; background-color: #808; color: white; padding: 5px; margin: 10px; border-radius: 5px; font-size: 12pt;"><b>IMPLEMENT YOUR ANSWER IN THE FOLLOWING CODE CELL</b></span>

<!-- grade_id: cell-6khz054xcj9yz8s3, points: 0, grade: false, solution: true -->
```python
def dna_reverse( sequence ):
    return sequence[::-1]
     
```

### <span id="Q1p3" style="background-color: #515; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">TEST 1/2 (20 points)</span><a href="#QTOP" style="background-color: #aaa; color: white; padding: 5px; padding-bottom: 5px; border-radius: 2px; text-align: center; border-radius: 3px; margin-left: 10px; text-decoration: none;">⇡</a>

<!-- grade_id: cell-o6an411594b2gfk8, points: 20, grade: true, solution: false -->
```python
sequence = 'ATAACGTGGATGCTA'

student_answer = dna_reverse( sequence )
correct_answer = 'ATCGTAGGTGCAATA'

assert student_answer == correct_answer
print( "+++ CORRECT!" )

```

Output:
```
+++ CORRECT!
```

### <span id="Q1p3" style="background-color: #515; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">TEST 2/2 (20 points)</span><a href="#QTOP" style="background-color: #aaa; color: white; padding: 5px; padding-bottom: 5px; border-radius: 2px; text-align: center; border-radius: 3px; margin-left: 10px; text-decoration: none;">⇡</a>

<!-- grade_id: cell-0dce0df4be182d50, points: 20, grade: true, solution: false -->
```python
sequence = 'GTCATGCTAAGTCTGA'

student_answer = dna_reverse( sequence )
correct_answer = 'AGTCTGAATCGTACTG'

assert student_answer == correct_answer
print( "+++ CORRECT!" )


```

Output:
```
+++ CORRECT!
```

<div style="border-bottom: 2px solid #888; margin-bottom: 10px;"></div>

<span style="display: inline-block; width: 100%; background-color: #dd6a00; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px; text-align: center; font-weight: bold;">END OF HOMEWORK</span>